<a href="https://colab.research.google.com/github/QuanPhamAnh/Coursera_Capstone/blob/main/Applied%20Data%20Science%20-%20UMICH/Introduction%20to%20DS%20in%20Python/Assignment_2_Pandas_1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Assignment 2
For this assignment you'll be looking at 2017 data on immunizations from the CDC. Your datafile for this assignment is in [assets/NISPUF17.csv](assets/NISPUF17.csv). A data users guide for this, which you'll need to map the variables in the data to the questions being asked, is available at [assets/NIS-PUF17-DUG.pdf](assets/NIS-PUF17-DUG.pdf). **Note: you may have to go to your Jupyter tree (click on the Coursera image) and navigate to the assignment 2 assets folder to see this PDF file).**

```python
import pandas as pd
import numpy as np
df = pd.read_csv('/content/NISPUF17.csv', index_col = 0)
```

```python
df.head(2)
```

Output:
```
SEQNUMC  SEQNUMHH  PDAT  ...  XVRCTY9  INS_STAT2_I  INS_BREAK_I
1   128521     12852     2  ...      NaN          NaN          NaN
2    10741      1074     2  ...      NaN          NaN          NaN

[2 rows x 453 columns]
```

```python
df.shape
```

Output:
```
(28465, 453)
```

## Question 1
Write a function called `proportion_of_education` which returns the proportion of children in the dataset who had a mother with the education levels equal to less than high school (<12), high school (12), more than high school but not a college graduate (>12) and college degree.

*This function should return a dictionary in the form of (use the correct numbers, do not round numbers):* 
```
    {"less than high school":0.2,
    "high school":0.4,
    "more than high school but not college":0.2,
    "college":0.2}
```


<!-- grade_id: cell-eea16d020eb52ae7, points: 0, grade: false, solution: true -->
```python
def proportion_of_education():
    # your code goes here

    educ = {}
    all = df['EDUC1'].count()
    less_than_high_school = len(df[df['EDUC1'] == 1])/all
    high_school = len(df[df['EDUC1'] == 2])/all
    more_than_high_school = len(df[df['EDUC1'] == 3])/all
    college = len(df[df['EDUC1'] == 4])/all

    educ['less than high school'] = less_than_high_school
    educ['high school'] = high_school
    educ['more than high school but not college'] = more_than_high_school
    educ['college'] = college

    return educ

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c0eeef201366f51c, points: 1, grade: true, solution: false -->
```python
assert type(proportion_of_education())==type({}), "You must return a dictionary."
assert len(proportion_of_education()) == 4, "You have not returned a dictionary with four items in it."
assert "less than high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "high school" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "more than high school but not college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."
assert "college" in proportion_of_education().keys(), "You have not returned a dictionary with the correct keys."

```

```python
proportion_of_education()
```

Output:
```
{'college': 0.47974705779026877,
 'high school': 0.172352011241876,
 'less than high school': 0.10202002459160373,
 'more than high school but not college': 0.24588090637625154}
```

## Question 2

Let's explore the relationship between being fed breastmilk as a child and getting a seasonal influenza vaccine from a healthcare provider. Return a tuple of the average number of influenza vaccines for those children we know received breastmilk as a child and those who know did not.

*This function should return a tuple in the form (use the correct numbers:*
```
(2.5, 0.1)
```

<!-- grade_id: cell-77f18c512324eabb, points: 0, grade: false, solution: true -->
```python
def average_influenza_doses():
    
    breastmilk = df[df['CBF_01'] == 1]['P_NUMFLU'].sum()/len(df[df['CBF_01'] == 1]['P_NUMFLU'].dropna())
    not_breastmilk = df[df['CBF_01'] == 2]['P_NUMFLU'].sum()/len(df[df['CBF_01'] == 2]['P_NUMFLU'].dropna())

    return (breastmilk, not_breastmilk)

    raise NotImplementedError()
```

<!-- grade_id: cell-54a3ba6cff31caa7, points: 1, grade: true, solution: false -->
```python
assert len(average_influenza_doses())==2, "Return two values in a tuple, the first for yes and the second for no."

```

```python
average_influenza_doses()
```

Output:
```
(1.8799187420058687, 1.5963945918878317)
```

## Question 3
It would be interesting to see if there is any evidence of a link between vaccine effectiveness and sex of the child. Calculate the ratio of the number of children who contracted chickenpox but were vaccinated against it (at least one varicella dose) versus those who were vaccinated but did not contract chicken pox. Return results by sex. 

*This function should return a dictionary in the form of (use the correct numbers):* 
```
    {"male":0.2,
    "female":0.4}
```

Note: To aid in verification, the `chickenpox_by_sex()['female']` value the autograder is looking for starts with the digits `0.0077`.

<!-- grade_id: cell-a0a9e6fe67698006, points: 0, grade: false, solution: true -->
```python
def chickenpox_by_sex():
    # YOUR CODE HERE

    male = len(df[(df['P_NUMVRC'].gt(0)) & (df['HAD_CPOX'] == 1) & (df['SEX'] == 1)])/\
    len(df[(df['P_NUMVRC'].gt(0)) & (df['HAD_CPOX'] == 2) & (df['SEX'] == 1)]) 
    female = len(df[(df['P_NUMVRC'].gt(0)) & (df['HAD_CPOX'] == 1) & (df['SEX'] == 2)])/\
    len(df[(df['P_NUMVRC'].gt(0)) & (df['HAD_CPOX'] == 2) & (df['SEX'] == 2)])

    return {'male' : male, 'female' : female}

    raise NotImplementedError()
```

<!-- grade_id: cell-c4f1714db100c865, points: 1, grade: true, solution: false -->
```python
assert len(chickenpox_by_sex())==2, "Return a dictionary with two items, the first for males and the second for females."

```

```python
chickenpox_by_sex()
```

Output:
```
{'female': 0.0077918259335489565, 'male': 0.009675583380762664}
```

## Question 4
A correlation is a statistical relationship between two variables. If we wanted to know if vaccines work, we might look at the correlation between the use of the vaccine and whether it results in prevention of the infection or disease [1]. In this question, you are to see if there is a correlation between having had the chicken pox and the number of chickenpox vaccine doses given (varicella).

Some notes on interpreting the answer. The `had_chickenpox_column` is either `1` (for yes) or `2` (for no), and the `num_chickenpox_vaccine_column` is the number of doses a child has been given of the varicella vaccine. A positive correlation (e.g., `corr > 0`) means that an increase in `had_chickenpox_column` (which means more no’s) would also increase the values of `num_chickenpox_vaccine_column` (which means more doses of vaccine). If there is a negative correlation (e.g., `corr < 0`), it indicates that having had chickenpox is related to an increase in the number of vaccine doses.

Also, `pval` is the probability that we observe a correlation between `had_chickenpox_column` and `num_chickenpox_vaccine_column` which is greater than or equal to a particular value occurred by chance. A small `pval` means that the observed correlation is highly unlikely to occur by chance. In this case, `pval` should be very small (will end in `e-18` indicating a very small number).

[1] This isn’t really the full picture, since we are not looking at when the dose was given. It’s possible that children had chickenpox and then their parents went to get them the vaccine. Does this dataset have the data we would need to investigate the timing of the dose?

<!-- grade_id: cell-8afff07f564cf79a, points: 0, grade: false, solution: true -->
```python
def corr_chickenpox():
    import scipy.stats as stats
    import numpy as np
    import pandas as pd
    
    # this is just an example dataframe
    #df=pd.DataFrame({"had_chickenpox_column":np.random.randint(1,3,size=(100)),
    #              "num_chickenpox_vaccine_column":np.random.randint(0,6,size=(100))})
    cpox_df = df[df['HAD_CPOX'].lt(3)].loc[:,['HAD_CPOX','P_NUMVRC']].dropna()
    cpox_df.columns=['had_chickenpox_column','num_chickenpox_vaccine_column']    

    # here is some stub code to actually run the correlation
    corr, pval=stats.pearsonr(cpox_df["had_chickenpox_column"], cpox_df["num_chickenpox_vaccine_column"])
    
    # just return the correlation
    
    return corr

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-73408733533a29a5, points: 1, grade: true, solution: false -->
```python
assert -1<=corr_chickenpox()<=1, "You must return a float number between -1.0 and 1.0."

```

```python
corr_chickenpox()
```

Output:
```
0.07044873460148046
```

```python

```